Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

**Provide your name and any collaborators below:**

<!-- grade_id: name-collabs, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

```scala
//mcyinsz
```

# Lab 1 Intro to Scala and Chisel
> Labs will be due each week before the homeworks. They are not intended take a significant amount of time but rather to provide examples/practice on specific and isolated features in the language. Labs are autograded so you can get quick feedback.

### Import the necessary Chisel dependencies. 
> There will be a cell like this in every lab. Make sure you run it before proceeding to bring the Chisel Library into the Jupyter Notebook scope!

```scala
interp.load.module(os.Path(s"${System.getProperty("user.dir")}/resource/chisel_deps.sc"))
```

```scala
import chisel3._
import chisel3.util._
import chiseltest._
import chiseltest.RawTester.test
```

Output:
```
[32mimport [39m[36mchisel3._
[39m
[32mimport [39m[36mchisel3.util._
[39m
[32mimport [39m[36mchiseltest._
[39m
[32mimport [39m[36mchiseltest.RawTester.test[39m
```

## Problem 1 (2 pts) - Scala conditionals
> Practice with Scala's if/else. Return `"heads"` if `flip` is `true`, else return `"tails"`

<!-- grade_id: ScalaCondPractice, points: 0, grade: false, solution: true -->
```scala
def scalaCondPractice(flip: Boolean): String = {
    // YOUR CODE HERE
    if (flip){
        "heads"
    }
    else
    {
        "tails"
    }
    
}
```

Output:
```
defined [32mfunction[39m [36mscalaCondPractice[39m
```

<!-- grade_id: scp-test, points: 2, grade: true, solution: false -->
```scala
assert(scalaCondPractice(true) == "heads")
assert(scalaCondPractice(false) == "tails")
```

## Problem 2 (3 pts) - Writing a Chisel Module's IO
> Fill in the IO in the module such that it takes two 4-bit `UInt`s as input and returns a 5-bit sum as output. 

<!-- grade_id: addTwo, points: 0, grade: false, solution: true -->
```scala
class AddTwo extends Module {
    // val io = ???
    // YOUR CODE HERE
    val io = IO(
        new Bundle {
            val in1=Input(UInt(4.W)) //指定n.W为位宽
            val in2=Input(UInt(4.W))
            val out=Output(UInt(5.W))
        }
    )

    
    io.out := io.in1 +& io.in2
}
```

Output:
```
defined [32mclass[39m [36mAddTwo[39m
```

<!-- grade_id: testAddTwo, points: 3, grade: true, solution: false -->
```scala
def testAddTwo: Boolean = {
    test(new AddTwo) { c =>
        for (i <- 0 until 16) {
            for (j <- 0 until 16) {
                c.io.in1.poke(i.U)
                c.io.in2.poke(j.U)
                c.io.out.expect((i+j).U)
            }
        }
    }
    true
}

assert(testAddTwo)
```

Output:
```
defined [32mfunction[39m [36mtestAddTwo[39m
```

## Problem 3 (3 pts) - Combinational Logic
> Assign the boolean expression: `(a AND b) OR (NOT c)` to the module's output. 

<!-- grade_id: CombLogic, points: 0, grade: false, solution: true -->
```scala
class CombLogic extends Module {
    val io = IO(new Bundle {
        val a   = Input(Bool())
        val b   = Input(Bool())
        val c   = Input(Bool())
        val out = Output(Bool())
    })
    
    // YOUR CODE HERE
    io.out := (io.a && io.b)||(!io.c)
    
    // We can print state like this everytime `step()` is called in our test
    printf(p"a: ${io.a}, b: ${io.b}, c: ${io.c}, out: ${io.out}\n")
}
```

Output:
```
defined [32mclass[39m [36mCombLogic[39m
```

```scala
visualize(() => new CombLogic)
```

## Problem 4 (4 pts) - Combinational Logic Test
> Write your own test that tests `CombLogic` exhaustively for all input values `a, b, and c`. The module should return `true` if and only if all calls to `dut.io.out.expect(...)` succeed.

<!-- grade_id: testCombLogic, points: 0, grade: false, solution: true -->
```scala
def testCombLogic: Boolean = {
    test(new CombLogic) { dut =>
        
        // YOUR CODE HERE
        dut.io.a.poke(1.B)
        dut.io.b.poke(1.B)
        dut.io.c.poke(1.B)
        dut.io.out.expect(1.B)

        dut.io.a.poke(0.B)
        dut.io.b.poke(1.B)
        dut.io.c.poke(1.B)
        dut.io.out.expect(0.B)

        dut.io.a.poke(0.B)
        dut.io.b.poke(0.B)
        dut.io.c.poke(1.B)
        dut.io.out.expect(0.B)
    }
    true
}
```

Output:
```
defined [32mfunction[39m [36mtestCombLogic[39m
```

<!-- grade_id: testCombLogic-test, points: 4, grade: true, solution: false -->
```scala
assert(testCombLogic)

```

## Problem 5 (3 pts) - Scala Conditional in Chisel modules
> At hardware ellaboration time, we can use Scala conditionals to change which hardware is created within a module. Implement the module such that if the `useAnd` argument is `true`, the generated hardware produces `a && b`, and otherwise produces `a || b`. The generated hardware should contain only `AND` logic or `OR` logic, but not both.

<!-- grade_id: AndOrCompileTime, points: 0, grade: false, solution: true -->
```scala
class AndOrGenerationTime(useAnd: Boolean) extends Module {
    val io = IO(new Bundle {
        val a   = Input(Bool())
        val b   = Input(Bool())
        val out = Output(Bool())
    })
    
    // YOUR CODE HERE
    if (useAnd){
        io.out:=io.a&&io.b
    }
    else{
        io.out:=io.a||io.b
    }
}
```

Output:
```
defined [32mclass[39m [36mAndOrGenerationTime[39m
```

```scala
visualize(() => new AndOrGenerationTime(false))
```

<!-- grade_id: testAndOrCompileTime, points: 3, grade: true, solution: false -->
```scala
def testAndOrGenerationTime(useAnd: Boolean): Boolean = {
    test(new AndOrGenerationTime(useAnd)) { dut =>
        for (a <- Seq(true, false)) {
            for (b <- Seq(true, false)) {
                dut.io.a.poke(a.B)
                dut.io.b.poke(b.B)
                if (useAnd) dut.io.out.expect((a && b).B)
                else        dut.io.out.expect((a || b).B)
            }
        }
    }
    true
}
assert(testAndOrGenerationTime(useAnd = true))
assert(testAndOrGenerationTime(useAnd = false))
```

Output:
```
defined [32mfunction[39m [36mtestAndOrGenerationTime[39m
```

## Problem 6 (3 pts) - Chisel Conditional in Chisel modules
> The generated hardware can use conditionals (i.e `Mux` or `when/elsewhen/otherwise`) to select signals. In this exercise, `useAnd` is an `Input` to the module. If `useAnd` is `true`, then the output `out` should be `a && b`, otherwise `a || b`. In this example both the logic for `a && b` and `a || b` hardware will be generated. You may use either a Chisel `Mux` or the Chisel `when` statements.

<!-- grade_id: AndOrRunTime, points: 0, grade: false, solution: true -->
```scala
class AndOrRunTime extends Module {
    val io = IO(new Bundle {
        val a      = Input(Bool())
        val b      = Input(Bool())
        val useAnd = Input(Bool())
        val out    = Output(Bool())
    })
    
    // YOUR CODE HERE
    io.out:=(io.a&&io.b&&io.useAnd)||((io.a||io.b)&&(!io.useAnd))
    
}
```

Output:
```
defined [32mclass[39m [36mAndOrRunTime[39m
```

```scala
visualize(()=>new AndOrRunTime)
```

<!-- grade_id: testAndOrRunTime, points: 3, grade: true, solution: false -->
```scala
def testAndOrRunTime: Boolean = {
    test(new AndOrRunTime) { dut =>
        for (a <- Seq(true, false)) {
            for (b <- Seq(true, false)) {
                for (useAnd <- Seq(true, false)) {
                    dut.io.a.poke(a.B)
                    dut.io.b.poke(b.B)
                    dut.io.useAnd.poke(useAnd.B)
                    if (useAnd) dut.io.out.expect((a && b).B)
                    else        dut.io.out.expect((a || b).B)
                }
            }
        }
    }
    true
}
assert(testAndOrRunTime)
```

Output:
```
defined [32mfunction[39m [36mtestAndOrRunTime[39m
```

## Problem 7 (2 pts) - Last connect semantics
> When connecting Chisel components, the last connection made is the one that "wins" (exists in the generated hardware). In the module below, the default output is `5.U` because `out` is connected to `5.U` after `4.U`. Use a `when` statement to conditionally connect `10.U` to the output when the input `update` is set high, or keep the default connection when `update` is set low.

<!-- grade_id: LastConnect, points: 0, grade: false, solution: true -->
```scala
class LastConnect extends Module {
    val io = IO(new Bundle {
        val update = Input(Bool())
        val out    = Output(UInt())
    })
    
    io.out := 4.U
    io.out := 5.U
    // YOUR CODE HERE
    when(io.update){
        io.out:=10.U
    }.otherwise{

    }
    
}
```

Output:
```
defined [32mclass[39m [36mLastConnect[39m
```

```scala
visualize(()=>new LastConnect)
```

<!-- grade_id: testLastConnect, points: 2, grade: true, solution: false -->
```scala
def testLastConnect: Boolean = {
    test(new LastConnect) { dut =>
        dut.io.update.poke(true.B)
        dut.io.out.expect(10.U)
        
        dut.io.update.poke(false.B)
        dut.io.out.expect(5.U)
    }
    true
}
assert(testLastConnect)
```

Output:
```
defined [32mfunction[39m [36mtestLastConnect[39m
```

## Problem 8 (8 pts) - Simple ReLU
> Let's put together some of these techniques to build a more complicated module. A ReLU or rectified linear unit is a function used in ML. (https://en.wikipedia.org/wiki/Rectifier_(neural_networks))

> To combine everything we've learned so far we will slightly modify the function to saturate at a parameterized upper-bound of our choosing. The module will compute this function: `f(x, upper_bound) = max(0, min(x, upperBound))`.

> Here is an example where we parameterize `upperBound = 3`. Note the input `x` and output `y` will be of type SInt.

<img src="images/relu.png" style="width:60%;">

<!-- grade_id: relu, points: 0, grade: false, solution: true -->
```scala
class ReLU(upperBound: Int) extends Module {
    val io = IO(new Bundle {
        val x = Input(SInt(5.W))
        val y = Output(SInt(5.W))
    })
    // YOUR CODE HERE
    when(io.x>upperBound.S){
        io.y:=upperBound.S
    }.elsewhen(io.x<0.S){
        io.y:=0.S
    }.otherwise{
        io.y:=io.x
    }
    
}
```

Output:
```
defined [32mclass[39m [36mReLU[39m
```

```scala
visualize (()=>new ReLU(2))
```

### Testing ReLU
> Write your own test for `ReLU` that tests `x` at input values `-1, 0, 1, 15`. The test (`testReLU`) is parameterized by `upperBound`, and you can assume `upperBound` is non-negative. The module should return `true` if and only if all calls to `dut.io.y.expect(...)` succeed.

<!-- grade_id: testReLu, points: 0, grade: false, solution: true -->
```scala
def testReLU(upperBound: Int): Boolean = {
    require(upperBound >= 0)
    test(new ReLU(upperBound)) { dut =>
        
    // YOUR CODE HERE
    dut.io.x.poke(-1)
    dut.io.y.expect(0)

    dut.io.x.poke(0)
    dut.io.y.expect(0)

    dut.io.x.poke(1)
    if(upperBound>1){
        dut.io.y.expect(1)
    }
    else{
        dut.io.y.expect(upperBound)
    }

    dut.io.x.poke(15)
    if(upperBound>15){
        dut.io.y.expect(15)
    }
    else{
        dut.io.y.expect(upperBound)
    }

    }
    true
}

```

Output:
```
defined [32mfunction[39m [36mtestReLU[39m
```

<!-- grade_id: testRelu, points: 8, grade: true, solution: false -->
```scala
for(upperBound <- 0 until 16) {
    println(s"Testing ReLu, upperBound=$upperBound")
    assert(testReLU(upperBound))
}
```

Output:
```
Testing ReLu, upperBound=0
Testing ReLu, upperBound=1
Testing ReLu, upperBound=2
Testing ReLu, upperBound=3
Testing ReLu, upperBound=4
Testing ReLu, upperBound=5
Testing ReLu, upperBound=6
Testing ReLu, upperBound=7
Testing ReLu, upperBound=8
Testing ReLu, upperBound=9
Testing ReLu, upperBound=10
Testing ReLu, upperBound=11
Testing ReLu, upperBound=12
Testing ReLu, upperBound=13
Testing ReLu, upperBound=14
Testing ReLu, upperBound=15
```